# Artificial Intelligence for Robotics 04

## General Information:
Please do not add or delete any cells. Answers belong into the corresponding cells (below the question). If a function is given (either as a signature or a full function), you should not change the name, arguments or return value of the function.

If you encounter empty cells underneath the answer that can not be edited, please ignore them, they are for testing purposes.

When editing an assignment there can be the case that there are variables in the kernel. To make sure your assignment works, please restart the kernel and run all cells before submitting (e.g. via Kernel -> Restart & Run All).

Code cells where you are supposed to give your answer often include the line `raise NotImplementedError`. This makes it easier to automatically grade answers. If you edit the cell please outcomment or delete this line.

## Submission:
Please submit your notebook via the web interface (in the main view -> Assignments -> Submit). The assignments are **due on Monday at 20:00**.

## Group Work:
Please enter your UID (your username here) and those of your group partners into the next cell. We apply plagiarism checking, so do not submit others solutions! If an assignment has a copied solution, the task will be graded with 0 points for all people with the same solution.

## Questions about the Assignment:
If you have questions about the assignment please post them in the LEA forum before the deadline. Don't wait until the last day to post questions!

### Please add the usernames of all your team members in the manner member1, member2 in next cell (example given below)

member1 = 'example'

member2 = 'example2'

#### If you are not working in a group, then please add member2 as none2s

<!-- grade_id: cell-a9c2f4587c671f98, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
member1 = 'ahoosh2s'
member2 = 'anuhel2s'
#raise NotImplementedError()
```

<!-- grade_id: cell-e27ced94e387fe18, points: 0, grade: true, solution: false -->
```python
# Execute this cell to make sure you correctly filled in the usernames of the team members

def group_name_test():
    for member_id in [member1, member2]:
        assert isinstance(member_id, str), "Please give your member id as a string."
        assert len(member_id) > 0, "You need to fill in the member id for both members"
        assert member_id.endswith("2s"), "The member id should end with 2s (Your JupyterHub username)"

group_name_test() 
print("All tests passed!")
```

Output:
```
All tests passed!
```

# Task 1

**[18 Point(s)]**


# A* Theory [20 Points]

During the lecture, you have learned about the A* search algorithm and its propoerties. Please answere the following questions short and consice and be as precise as possible.


## Task 1.1

**[2 Point(s)]**

### Completeness [2 Points]

How is the *completeness* of an algorithm defined? Is A* complete?

<!-- grade_id: Astar_theory_A, points: 2, grade: true, solution: true -->
A* is complete if the heuristic function is admissible, meaning that the cost considered by the heuristic function is never more that the actual cost. Additionally, completeness of A* relies upon the search space which should finite. 

There are times when the search space is infinite, and in such a suituation only if the heuristic fucntion is both admissible and consistent can it claim to be complete.

## Task 1.2

**[2 Point(s)]**

### Soundness [2 Points]

When is an algorithm be considered *sound*? Is A* sound?

<!-- grade_id: Astar_theory_B, points: 2, grade: true, solution: true -->
A sound algorithm never allows in a wrong answer into the set of right answers, so the results of such an algorithm can always be trusted. However, in some cases a sound algorithm may not provide an answer. For this reason, if the heuristic function is admissible meaning that it never overestimates the cost of paths to the goal, and along with it is consistent which guarantees that no other path will reuslt in a more optimal traveral cost, it is definitely sound. In other words, if the hueristic fcuntion is both admissible and consistent, A* renders the optimal path to the goal in terms of cost. 

## Task 1.3

**[2 Point(s)]**

### Termination [2 Points]

When does A* end the search process?

<!-- grade_id: Astar_theory_C, points: 2, grade: true, solution: true -->
A* generally terminates when it finds the optimal most cost-efficient path from the initial state to the goal. Upon termination, it guarantees that no other path is more optimal than the one found by the algorithm, assuming that the heuristic fucntion fulfil the requirements of admissibility and consistency.  

## Task 1.4

**[5 Point(s)]**

### Optimality [5 Points]

What does it mean when we say an algorithm is *optimal*? What are the criteria that have to be met in order for A* to be *optimal*? Provide an argument or proof on why this is required.

<!-- grade_id: Astar_theory_D, points: 5, grade: true, solution: true -->
Optimality is an attribute of an algorithm which ensures that a solution provided by it is always the lowest in terms of cost. For A* to be optimal, the huersitic fuction should be both admissible and consistent.

Admissibility:

The huersitic should never overestimate the cost of traversal from a given node - n to the goal

h(n) <= h*(n)

wheren h*(n) is the actual cost of node -n to the goal

If the heuristic fucntion everestimates the actual cost and, in other words, exceeds it, then path found by the heuristic fucntion could be even greater than the optimal path to the goal, meaning that a non-optimal path is allowed into consideration.

Consistency:

If we consider n' as a successor of n and the cost of n to n' is shown as c(n,n') the heuristic fucntion of h(n) should always fulfil the following relation

h(n) <= c(n,n') + h(n')

This ensures that the sub path selected by A* is always optimal, too and there is no risk of losing sight of a lowest cost sub path from every successor. This also shows that the heuristic fucntion is never descending along the path from intial state to the goal. 

## Task 1.5

**[5 Point(s)]**

### Time Complexity [5 Points]

What is the time complexity of A*? Explain, or derive mathematically! Just writing $O(...)$ is not sufficient!

<!-- grade_id: Astar_theory_E, points: 5, grade: true, solution: true -->
There are different situations where Big O notation is defined for the A* search algorithm.

Let's considered the following relation:

f(n) = h(n) + g(n)

if h(n)=0 so the the actual cost of node n is considered, resulting in A* functioning like Uniform Cost Search in which the brancing factor, b, and the shallowest possible depth to the goal, d, play a role in determining the worst time complexity.

In such a case time complexity of A* is exponential:

O(b to the power of d)

Let's imagine that the heuristic function is perfectly equal to the actual cost:

h(n) = h*(n)

In this case, only the nodes on the optimal path will be taken into consideration, avoiding the other successors. Therefore, the shollowest depth to the solution is of significance in the caculation of Big O notation, so 

O(d) is the worst case scenario in terms of search time.

Last but not least is when the heuristic function is admissible and consistent:

In such a case, C* = h(n) + g(n) 

C* is the optimal bath to the goal, so all the nodes where f(n) <= C* are exanded, resulting in an exponential search space.

Here we need to consider the minimum step cost to come up with an estimate for the number of nodes exanded at each depth level

let's say that E is the minimum actual cost from between to candidate nodes, hence the number of each number nodes exanded at each level could C*/E

Then as  it a tree seacrh space then the Big O is:

O(braching factor to the power of depth)

and then it is

O (b ^ C*/E)

## Task 1.6

**[2 Point(s)]**

### Consistency of a Heuristic [2 Points]

How is a consistent heuristic defined? Why is it a critical property in the context of the A* search algorithm?

<!-- grade_id: Astar_theory_F, points: 2, grade: true, solution: true -->
Consistency:

If we consider n' as a successor of n and the cost of n to n' is shown as c(n,n') the heuristic fucntion of h(n) should always fulfil the following relation

h(n) <= c(n,n') + h(n')

This ensures that the sub path selected by A* is always optimal, too and there is no risk of losing sight of a lowest cost sub path from every successor. This also shows that the heuristic fucntion is never descending along the path from intial state to the goal.

Therefore, this ensures that for every child of the current path the most optimal sub-path will be selected by the A* algorithm, provided that the heuristic function is consistent, thereby satisfying the above criterion.

# Task 2

**[50 Point(s)]**

# Applying Informed Search [60 Points]

Informed methods help us gain information about solving a problem through its current state space. This keeps a program from blundering around blindly guessing. Informed search strategies make use of information about the path of moves we took to get where we are currently by using an *evaluation function*. This evaluation function often makes use of a *heuristic function*, that provides an estimate of the distance of a node to the goal in a relaxed version of the problem.

# A Real World Problem

Consider a robot that works in a warehouse. The robot can carry packages on top of itself on a set of omnidirectional conveyors. It can carry a maximum of 8 packages in a 3x3 grid, with one spot left open. There is a slot from which the robot can dispense packages onto shelves. The robot must deposit the packages in order, so they must be lined up in a certain order near the slot. However, workers place the packages onto the robot randomly. Therefore, the robot must first sort the packages before depositing them. It can use the free space and the conveyors to shift the packages horizontally and vertically. 

## This becomes what is known as the 8 puzzle.

## Task 2.1

**[15 Point(s)]**

## The 8 Puzzle [10 Points]

An 8 puzzle is a simple game consisting of a 3 x 3 grid (containing 9 squares). One of the squares is empty.
The object is to move the squares around into different positions and having the numbers displayed in the "goal state". The squares can obviously only move along the horizontal and vertical axis and only into the **empty tile**. Therefore the problem can also be viewed as the empty tile "trading" positions with one of its neighbouring tiles.

The image below depicts the goal state of an "3 x 3" 8 puzzle.
<img src="Informed_Search_files/img/8_puzzle.png" alt="8 Puzzle" width="200"/>

The complexity of possible moves toward the final solution in a game like this is large. With an uninformed search strategy, finding a solution to this problem can take a lot of time. The problem is now about finding the shortest path to the goal state within reasonable time, which we will do in this assignment using A* and Greedy Search.

### Task
- We will implement the 8 puzzle game as the class `Puzzle`
- Define methods for moving a tile around in the puzzle
- Define a mehod for generating all possible moves from one particular configuration.
- You are already provided with a class and function templates below. Please complete the implementation.

<!-- grade_id: Astar_implementation_A, points: 0, grade: false, solution: true -->
```python
from typing import List, Tuple
from utils import print_puzzle

class Puzzle:

    def __init__(self, init_state: List[int], puzzle_type: int = 8):
        """
        Initialize a new 8-puzzle board
        :param init_state: Initial configuration of the board
        """
        self.init_state = init_state
        self.puzzle_type = puzzle_type
        self.goal_state = [i for i in range (1, puzzle_type + 1)] + [0]
        self.closed =[]
        # YOUR CODE HERE
        #raise NotImplementedError()

    def goal_test(self, state: List[int]):
        """Test if goal state is reached
        :param state: board configuration to check
        :return: true if the passed configuration is equal to goal configuration
        """
        # Check if each element in the current state matches the goal configuration
        return state == self.goal_state
        # YOUR CODE HERE
        #raise NotImplementedError()
        
    def add_to_explored():
        
        self.closed.append(state)

    def is_explored(self, state: List[int]):
        """Check if a particular board configuration has already been explored
        :param state: board configuration to check
        :return: true if a particular configuration has already been explored
        """
        return state in self.closed
        # YOUR CODE HERE
        #raise NotImplementedError()
    
    def reset(self):
        # YOUR CODE HERE
        self.closed.clear()
        self.init_state = [i for i in range (1, self.puzzle_type + 1)] + [0]
        #raise NotImplementedError()


def move_left(position: int) -> int:
    """Move one position left in 8 puzzle if possible
    :param position: current position of the 0 tile
    :return: new position of the 0 tile after moving to the left
    """
    # YOUR CODE HERE
    if position % 3 > 0:
        return position -1
    return None
    #raise NotImplementedError()


def move_right(position: int) -> int:
    """Move one position right in 8 puzzle if possible
    :param position: current position of the 0 tile
    :return: new position of the 0 tile after moving to the right
    """
    # YOUR CODE HERE
    if position % 3 < 2:
        return position + 1
    return None
    #raise NotImplementedError()


def move_up(position: int) -> int:
    """Move one position up in 8 puzzle if possible
    :param position: current position of the 0 tile
    :return: new position of the 0 tile after moving upwards
    """
    # YOUR CODE HERE
    if position > 2:
        return position -3
    return None
    #raise NotImplementedError()


def move_down(position: int):
    """Move one position down in 8 puzzle if possible
    :param position: current position of the 0 tile.
    :return: new position of the 0 tile after moving downwards
    """
    # YOUR CODE HERE
    if position < 6:
        return position + 3
    return None
    #raise NotImplementedError()


def get_possible_moves(state: List[int]) -> List[List[int]]:
    """Check whether a move is possible in left, right, up, down direction and store it.
    :param state: current configuration of the puzzle as one dimensional list
    :return: list containing the new configurations after applying all possible moves
    """
    # YOUR CODE HERE
    moves = []
    
    position = state.index(0)
    
    for move_func in (move_left, move_right, move_up, move_down):
        
        new_position = move_func (position)
        if new_position is not None:
            new_state = state[:]
            
            new_state [position], new_state[new_position] = new_state [new_position], new_state [position]
            
            moves.append (new_state)
    
    return moves
    #raise NotImplementedError()

# YOUR CODE HERE
#raise NotImplementedError()
```

<!-- grade_id: test_Astar_implementation_A0, points: 15, grade: true, solution: false -->
```python

```

## Task 2.2

**[5 Point(s)]**

## Heuristics [5 Points]

In this part, you are tasked to implement two different admissible heuristics for the 8 puzzle: The [Hamming distance](https://en.wikipedia.org/wiki/Hamming_distance) and the [Manhattan Distance](https://en.wikipedia.org/wiki/Taxicab_geometry).

**Hamming Distance**  
Let $A$ be an alphabet of symbols and $C$ a subset of $A^n$, the set of worlds of length $n$ over $A$. Let $u = (u_1, ..., u_n)$ and $v = (v_1, ...,v_n)$. The Hamming distance is deinfed as the number of places in which $u$ and $v$ differ. That is:
$$h(n) = \{i:u_i \neq v_i,i=1,...,n\}$$
See also: [Encyclopedia of Math](https://encyclopediaofmath.org/wiki/Hamming_distance)

It is clear that this heuristic is admissible since the total number of moves to order the tiles correctly is at least the number of misplaced tiles (each tile not in place must be moved at least once). The cost (number of moves) to the goal (an ordered puzzle) is at least the Hamming distance of the puzzle.

**Manhattan Distance**  
The Manhattan distance between two points $x = (x_1, ..., x_n)$ and $y = (y_1, ..., y_n)$ in $n$-dimensional space is the sum of the distances in each dimension.
$$d(x,y) = \sum^n_{i=1} \vert x_i - y_i \vert$$
See also: [Encyclopedia of Machine Learning](https://link.springer.com/referenceworkentry/10.1007%2F978-0-387-30164-8_506)

The Manhattan distance is an admissible heuristic because every tile will have to be moved at least the number of spots in between itself and its correct position.

### Task
- You are provided with the function definitions for the two heuristics below
- Please implement them according to the mathematical definition given above

<!-- grade_id: Astar_implementation_B, points: 0, grade: false, solution: true -->
```python
from math import sqrt
from typing import Tuple, List

def hamming_distance(state: List[int]) -> int:
    """Calculate the Hamming distance for a particular configuration
    :param state: current configuration of the puzzle
    :return: the number of misplaced tiles in the given configuration
    """
    hamming = 0
    
    for index, item in enumerate (state):
        
        if item != 0 and item != index + 1:
            
            hamming += 1
            
        elif item == 0 and item != len(state) -1:
            
            hamming += 1
            
    return hamming
    # YOUR CODE HERE
    #raise NotImplementedError()


def manhattan_distance(state: List[int]) -> int:
    """Function to calculate the manhattan distance for a
    particular configuration
    :param state: current configuration of the puzzle
    :return: the accumulated manhattan distance between each tile and its goal position in the given configuration
    """
    manhattan = 0
    reference = int (sqrt(len(state)))
    
    for index, item in enumerate (state):
        
            if item != 0:
                
                goal_position_x, goal_position_y = divmod(item-1, reference)
            else:
                goal_position_x, goal_position_y = reference - 1, reference -1

            current_position_x, current_position_y = divmod(index, reference)

            manhattan += abs(goal_position_x - current_position_x) + abs (goal_position_y - current_position_y)
            
    return manhattan
    # YOUR CODE HERE
    #raise NotImplementedError()
```

```python
# Run this test cell in order to get an indication of whether your implementation is working as expected.
assert hamming_distance([0,1,2,3,4,5,8,6,7]) == 9
assert manhattan_distance([1,2,3,4,0,5,6,8,7]) == 8
```

<!-- grade_id: test_Astar_implementation_B0, points: 5, grade: true, solution: false -->
```python

```

## Task 2.3

**[10 Point(s)]**

## A* Search Algorithm [10 Points]

Now you will implement the A* search algorithm. To do that, you can consider the pseudocode from the lecture

```
function A*(problem) returns a solution, or failure
    closed ← an empty set
    fringe ← a list containing Make-Node(Initial-State[problem])
    loop do
        if fringe is empty then return failure
        node ← Remove-Front(fringe)
        if Goal-Test[problem] applied to State(node) succeeds return node
            insert node into closed
        for each node n ∈ Expand(node, problem) do
            if there is a node m ∈ closed ∪ fringe such that
                STATE(m) = STATE(n) and f(m) ≤ f(n)
            then do nothing
            else
                insert n into fringe after the last node m such that f(m) ≤ f(n)
end
```
Some other references for A* pseudocode: https://mat.uab.cat/~alseda/MasterOpt/AStar-Algorithm.pdf
### Task
- You are provided with the function definition for the A* function
- Please implement A* according to the algorithm design discussed in the lectures

In order to get an estimation on whether your solution work as expected, there is a public test available in the test cell. Make sure that all dependent cells are executed first, otherwise some functions or variables might not exist and an error will occur.

<!-- grade_id: Astar_implementation_C, points: 0, grade: false, solution: true -->
```python
from typing import Callable, List, Tuple, Dict, Optional
import heapq

def astar_search(board: Puzzle, heuristic: Callable) -> Tuple[List[int], int]:
    """
    :param board: the 8-puzzle to solve
    :param heuristic: the heuristic function to use
    :return: an ordered list with the solution path and the number of total nodes expanded
    """
    fringe = []
    heapq.heappush(fringe, (0, board.init_state))
    source: Dict [Tuple[int], Optional[Tuple[int]]] = {}
    g_costs = {tuple (board.init_state): 0}
    explored = set()
    expanded_nodes = 0
    
    while fringe:
        
        _, current_state = heapq.heappop (fringe)
        expanded_nodes += 1
        
        if board.goal_test (current_state):
            return track_path (source, current_state), expanded_nodes
             
        
        explored.add(tuple(current_state))

        for child in get_possible_moves (current_state):
            child_tuple = tuple(child)
            tentative_g_cost = g_costs[tuple(current_state)] + 1

            if child_tuple in explored or (child_tuple in g_costs and tentative_g_cost >= g_costs[child_tuple]):
                continue
            
            source[child_tuple] = current_state
            g_costs[child_tuple] = tentative_g_cost
            
            f_cost = tentative_g_cost + heuristic (child)
            heapq.heappush (fringe, (f_cost,child))
        
    return [], expanded_nodes

def track_path(source: Dict[Tuple[int], Optional[Tuple[int]]], current: List[int]) -> List[int]:
    
    path = [current]
    while tuple(current) in source:
        current = source [tuple (current)]
        path.append(current)
    path.reverse()
    print (path)
    return path
    # YOUR CODE HERE
    #raise NotImplementedError()
```

```python
# Run this test cell in order to get an INDICATION of whether your implementation is working as expected.
import time
t0 = time.time()
p = Puzzle([0,1,2,3,4,5,8,6,7])
path, expanded_nodes = astar_search(p, manhattan_distance)
t1 = time.time()
assert len(path) == 23
# The execution time is limited to 10min, so please keep this in mind when implementing your solution!
assert t1-t0 <= 60
```

Output:
```
[[0, 1, 2, 3, 4, 5, 8, 6, 7], [1, 0, 2, 3, 4, 5, 8, 6, 7], [1, 4, 2, 3, 0, 5, 8, 6, 7], [1, 4, 2, 0, 3, 5, 8, 6, 7], [1, 4, 2, 8, 3, 5, 0, 6, 7], [1, 4, 2, 8, 3, 5, 6, 0, 7], [1, 4, 2, 8, 3, 5, 6, 7, 0], [1, 4, 2, 8, 3, 0, 6, 7, 5], [1, 4, 2, 8, 0, 3, 6, 7, 5], [1, 4, 2, 8, 7, 3, 6, 0, 5], [1, 4, 2, 8, 7, 3, 0, 6, 5], [1, 4, 2, 0, 7, 3, 8, 6, 5], [1, 4, 2, 7, 0, 3, 8, 6, 5], [1, 0, 2, 7, 4, 3, 8, 6, 5], [1, 2, 0, 7, 4, 3, 8, 6, 5], [1, 2, 3, 7, 4, 0, 8, 6, 5], [1, 2, 3, 7, 4, 5, 8, 6, 0], [1, 2, 3, 7, 4, 5, 8, 0, 6], [1, 2, 3, 7, 4, 5, 0, 8, 6], [1, 2, 3, 0, 4, 5, 7, 8, 6], [1, 2, 3, 4, 0, 5, 7, 8, 6], [1, 2, 3, 4, 5, 0, 7, 8, 6], [1, 2, 3, 4, 5, 6, 7, 8, 0]]
```

<!-- grade_id: test_Astar_implementation_C0, points: 10, grade: true, solution: false -->
```python

```

## Task 2.4

**[10 Point(s)]**

## Greedy Search [10 Points]

Now we will implement a *greedy search algorithm*. Recall from the lecture, that the evaluation function $f(n)$ is defined as $f(n) = h(n)$, with $h(n)$ being the heuristic function.

### Task
- You are provided with the definition of the greedy search function
- Implement a greedy search strategy by completing the function below

In order to get an estimation on whether your solution work as expected, there is a public test available in the test cell. Make sure that all dependent cells are executed first, otherwise some functions or variables might not exist and an error will occur.

<!-- grade_id: Astar_implementation_D, points: 0, grade: false, solution: true -->
```python
from typing import List, Tuple, Callable

def greedy_search(board: Puzzle, heuristic: Callable) -> Tuple[List[int], int]:
    """Implementation of the greedy search algorithm.
    :param board: the 8-puzzle to solve
    :param heuristic: the heuristic function to use
    :return: an ordered list with the solution path and the number of total nodes expanded 
    """
    # YOUR CODE HERE
    open_list = []
    heapq.heappush(open_list, (manhattan_distance(board.init_state), board.init_state))
    source: Dict [Tuple[int], Optional[Tuple[int]]] = {}
    # g_costs = {tuple (board.init_state): 0}
    explored = set()
    expanded_nodes = 0
    
    while open_list:
        
        hueristic_cost, current_state = heapq.heappop (open_list)
        expanded_nodes += 1
        
        if board.goal_test (current_state):
            return track_path (source, current_state), expanded_nodes
             
        
        explored.add(tuple(current_state))

        for child in get_possible_moves (current_state):
            child_tuple = tuple(child)

            if child_tuple in explored:
                continue
            
            source[child_tuple] = current_state
            heapq.heappush (open_list, (manhattan_distance(child), child))
        
    return [], expanded_nodes
    
    #raise NotImplementedError()

# YOUR CODE HERE
#raise NotImplementedError()
```

```python
# Run this test cell in order to get an INDICATION of whether your implementation is working as expected.
import time
t0 = time.time()
p = Puzzle([0,1,2,3,4,5,8,6,7])
path, expanded_nodes = greedy_search(p, manhattan_distance)
t1 = time.time()
# The execution time is limited to 10min, so please keep this in mind when implementing your solution!
assert t1-t0 <= 600
```

Output:
```
[[0, 1, 2, 3, 4, 5, 8, 6, 7], [1, 0, 2, 3, 4, 5, 8, 6, 7], [1, 2, 0, 3, 4, 5, 8, 6, 7], [1, 2, 5, 3, 4, 0, 8, 6, 7], [1, 2, 5, 3, 4, 7, 8, 6, 0], [1, 2, 5, 3, 4, 7, 8, 0, 6], [1, 2, 5, 3, 4, 7, 0, 8, 6], [1, 2, 5, 0, 4, 7, 3, 8, 6], [1, 2, 5, 4, 0, 7, 3, 8, 6], [1, 2, 5, 4, 7, 0, 3, 8, 6], [1, 2, 5, 4, 7, 6, 3, 8, 0], [1, 2, 5, 4, 7, 6, 3, 0, 8], [1, 2, 5, 4, 0, 6, 3, 7, 8], [1, 2, 5, 0, 4, 6, 3, 7, 8], [1, 2, 5, 3, 4, 6, 0, 7, 8], [1, 2, 5, 3, 4, 6, 7, 0, 8], [1, 2, 5, 3, 0, 6, 7, 4, 8], [1, 2, 5, 0, 3, 6, 7, 4, 8], [1, 2, 5, 7, 3, 6, 0, 4, 8], [1, 2, 5, 7, 3, 6, 4, 0, 8], [1, 2, 5, 7, 3, 6, 4, 8, 0], [1, 2, 5, 7, 3, 0, 4, 8, 6], [1, 2, 5, 7, 0, 3, 4, 8, 6], [1, 2, 5, 0, 7, 3, 4, 8, 6], [1, 2, 5, 4, 7, 3, 0, 8, 6], [1, 2, 5, 4, 7, 3, 8, 0, 6], [1, 2, 5, 4, 0, 3, 8, 7, 6], [1, 2, 5, 4, 3, 0, 8, 7, 6], [1, 2, 0, 4, 3, 5, 8, 7, 6], [1, 0, 2, 4, 3, 5, 8, 7, 6], [1, 3, 2, 4, 0, 5, 8, 7, 6], [1, 3, 2, 4, 5, 0, 8, 7, 6], [1, 3, 2, 4, 5, 6, 8, 7, 0], [1, 3, 2, 4, 5, 6, 8, 0, 7], [1, 3, 2, 4, 
```

<!-- grade_id: test_Astar_implementation_D0, points: 10, grade: true, solution: false -->
```python

```

## Task 2.5

**[5 Point(s)]**

## Evaluation A) [5 Points]

Now we want to compare the performance of the different solvers. For that, run every combination of `algorithm` x `heuristic` with the puzzle configurations
1. `[0, 1, 2, 3, 4, 5, 8, 6, 7]`
1. `[8, 7, 6, 5, 1, 4, 2, 0, 3]`
1. `[1, 5, 7, 3, 6, 2, 0, 4, 8]`

for each run, record the `expanded nodes`, `execution time` and the `path cost`.
```
[
  {
    'puzzle': 1,
    'algorithm': 'astar',
    'heuristic': 'manhattan',
    'expanded nodes': 0,
    'execution time': 0.0,
    'path cost': 0
  }
]
```

<!-- grade_id: Astar_implementation_E, points: 0, grade: false, solution: true -->
```python
import time
evaluation_results = []

puzzle_configurations = [[0, 1, 2, 3, 4, 5, 8, 6, 7], [8, 7, 6, 5, 1, 4, 2, 0, 3], [1, 5, 7, 3, 6, 2, 0, 4, 8]]
algorithms_to_test = {'astar': astar_search, 'greedy': greedy_search}
heuristics = {'manhattan': manhattan_distance, 'hamming': hamming_distance}

def trial_algorithms_heuristics():
    
    for configuration in puzzle_configurations:
        for algorithm_name, algorithm in algorithms_to_test.items():
            for heuristic_name, heuristic in heuristics.items():
                
                puzzle_to_test = Puzzle(configuration)
                
                start = time.time()
                path, expanded_nodes = algorithm(puzzle_to_test, heuristic)
                finish = time.time()
                
                path_cost = len(path) - 1 if path else 0

                
                result = {
                    'puzzle': configuration,
                    'algorithm': algorithm_name,
                    'heuristic': heuristic_name,
                    'expanded nodes': expanded_nodes,
                    'execution time': finish - start,
                    'path cost': path_cost
                }
                evaluation_results.append(result)
    
    for result in evaluation_results:
           print (result)

trial_algorithms_heuristics()
# YOUR CODE HERE
#raise NotImplementedError()
```

Output:
```
[[0, 1, 2, 3, 4, 5, 8, 6, 7], [1, 0, 2, 3, 4, 5, 8, 6, 7], [1, 4, 2, 3, 0, 5, 8, 6, 7], [1, 4, 2, 0, 3, 5, 8, 6, 7], [1, 4, 2, 8, 3, 5, 0, 6, 7], [1, 4, 2, 8, 3, 5, 6, 0, 7], [1, 4, 2, 8, 3, 5, 6, 7, 0], [1, 4, 2, 8, 3, 0, 6, 7, 5], [1, 4, 2, 8, 0, 3, 6, 7, 5], [1, 4, 2, 8, 7, 3, 6, 0, 5], [1, 4, 2, 8, 7, 3, 0, 6, 5], [1, 4, 2, 0, 7, 3, 8, 6, 5], [1, 4, 2, 7, 0, 3, 8, 6, 5], [1, 0, 2, 7, 4, 3, 8, 6, 5], [1, 2, 0, 7, 4, 3, 8, 6, 5], [1, 2, 3, 7, 4, 0, 8, 6, 5], [1, 2, 3, 7, 4, 5, 8, 6, 0], [1, 2, 3, 7, 4, 5, 8, 0, 6], [1, 2, 3, 7, 4, 5, 0, 8, 6], [1, 2, 3, 0, 4, 5, 7, 8, 6], [1, 2, 3, 4, 0, 5, 7, 8, 6], [1, 2, 3, 4, 5, 0, 7, 8, 6], [1, 2, 3, 4, 5, 6, 7, 8, 0]]
[[0, 1, 2, 3, 4, 5, 8, 6, 7], [1, 0, 2, 3, 4, 5, 8, 6, 7], [1, 2, 0, 3, 4, 5, 8, 6, 7], [1, 2, 5, 3, 4, 0, 8, 6, 7], [1, 2, 5, 3, 4, 7, 8, 6, 0], [1, 2, 5, 3, 4, 7, 8, 0, 6], [1, 2, 5, 3, 0, 7, 8, 4, 6], [1, 2, 5, 0, 3, 7, 8, 4, 6], [0, 2, 5, 1, 3, 7, 8, 4, 6], [2, 0, 5, 1, 3, 7, 8, 4, 6], [2, 3, 5, 1, 0, 7, 8, 4, 6], [2, 3, 5, 1,
```

<!-- grade_id: test_Astar_implementation_E0, points: 5, grade: true, solution: false -->
```python

```