# Artificial Intelligence for Robotics 03

## General Information:
Please do not add or delete any cells. Answers belong into the corresponding cells (below the question). If a function is given (either as a signature or a full function), you should not change the name, arguments or return value of the function.

If you encounter empty cells underneath the answer that can not be edited, please ignore them, they are for testing purposes.

When editing an assignment there can be the case that there are variables in the kernel. To make sure your assignment works, please restart the kernel and run all cells before submitting (e.g. via Kernel -> Restart & Run All).

Code cells where you are supposed to give your answer often include the line `raise NotImplementedError`. This makes it easier to automatically grade answers. If you edit the cell please outcomment or delete this line.

## Submission:
Please submit your notebook via the web interface (in the main view -> Assignments -> Submit). The assignments are **due on Monday at 20:00**.

## Group Work:
Please enter your UID (your username here) and those of your group partners into the next cell. We apply plagiarism checking, so do not submit others solutions! If an assignment has a copied solution, the task will be graded with 0 points for all people with the same solution.

## Questions about the Assignment:
If you have questions about the assignment please post them in the LEA forum before the deadline. Don't wait until the last day to post questions!

### Please add the usernames of all your team members in the manner member1, member2 in next cell (example given below)

member1 = 'example'

member2 = 'example2'

#### If you are not working in a group, then please add member2 as none2s

<!-- grade_id: cell-a9c2f4587c671f98, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
member1 = 'ahoosh2s'
member2 = 'anuhel2s'
#raise NotImplementedError()
```

<!-- grade_id: cell-e27ced94e387fe18, points: 0, grade: true, solution: false -->
```python
# Execute this cell to make sure you correctly filled in the usernames of the team members

def group_name_test():
    for member_id in [member1, member2]:
        assert isinstance(member_id, str), "Please give your member id as a string."
        assert len(member_id) > 0, "You need to fill in the member id for both members"
        assert member_id.endswith("2s"), "The member id should end with 2s (Your JupyterHub username)"

group_name_test() 
print("All tests passed!")
```

Output:
```
All tests passed!
```

# Task 1

**[36 Point(s)]**


# Introduction to Algorithms

You can get help from the excerpt of *Cormen’s* very nice book *Introduction to Algorithms*, available as a PDF ﬁle (`cormen_introalgorithms`) in LEA, to solve the following tasks.

## Task 1.1

**[10 Point(s)]**

### Task 1 [10 Points]

Why is a special notation needed to classify algorithms? Doesnt it suffice to merely measure the runtime in seconds? **Explain**.

<!-- grade_id: algorithms_introduction_A, points: 10, grade: true, solution: true -->
The runtime model is not a comprehensive explanation because it is influenced by software and hardware components. Additionally, runtime does not explain the relationship between the size of the input and the growth of runtime. For example, O(n) is a linear notation, while O(n^m) is exponential in nature, with its growth far exceeding O(n) for larger inputs. The difference between algorithms in terms of runtime cannot be fully explained through runtime results alone, so comparing algorithms in terms of time-efficiency and space-efficiency is impossible using, for example, a table of different runtime tests.

## Task 1.2

**[6 Point(s)]**

### Task 2 [6 Points]

Let $g : \mathbb{N} \rightarrow \mathbb{R}+$ be an arbitrary function. The set of functions $f : \mathbb{N} \rightarrow \mathbb{R}+$, which do not gorw faster than the function $g$ after a specific point $n_0$, is denoted as $O(g(n))$. More specifically:

$$
O(g(n) := f(n) \vert \exists c \in \mathbb{R}, \exists n_o \in \mathbb{N} : 0 \leq f(n) \leq c g(n) \forall n \geq n_0)
$$

**Prove** the following:
- $f(n) = 100n^2 \in O(n^2)$
- $f(n) = n^6 + 100n^5 \in O(n^6)$

<!-- grade_id: algorithms_introduction_B, points: 6, grade: true, solution: true -->
YOUR ANSWER HERE

## Task 1.3

**[3 Point(s)]**

### Task 3 [3 Points]

What is the asymptotic time complexity of the following python code? Assume that `ANY_CONST` is an arbitrary constant in your program, which has assigned a 2D array as its value.

```python
sum_ = 0
for i in range(0, J):
    for j in range(0, K):
        if arr[i][j] <= ANY_CONST:
            sum_ = sum_ + arr[i][j]
            print(sum_)
```

<!-- grade_id: algorithms_introduction_C, points: 3, grade: true, solution: true -->


## Task 1.4

**[15 Point(s)]**

### Task 4 [12 Points]

For each function $f(n)$ and time $t$ in the following table, determine the largest size $n$ of a problem that can be solved in time $t$, assuming that the algorithm to solve the problem takes $f(n)$ microseconds.

Show your **path to the solution** for:
- $log(n)$
- $2^n$
- $nlog(n)$.

| | 1 second | 1 minute | 1 hour | 1 day | 1 month | 1 year | 1 century |
|-|----------|----------|--------|-------|---------|--------|-------------|
| $\log_{10} n$ | | | | | | | |
| $\sqrt{n}$ | | | | | | | |
| $n$ | | | | | | | |
| $n \log_{10} n$ | | | | | | | |
| $n^2$ | | | | | | | |
| $n^3$ | | | | | | | |
| $2^n$ | | | | | | | |
| $n!$ | | | | | | | |

 > If you need to write any code, please do so in the code cell below.

<!-- grade_id: algorithms_introduction_D, points: 15, grade: true, solution: true -->
t=log(n)⟹n=10t

t=2**n ⟹ n=log2(t)

n*log(n) => is calculated numerically


| | 1 second | 1 minute | 1 hour | 1 day | 1 month | 1 year | 1 century |
|-|----------|----------|--------|-------|---------|--------|-------------|
| $\log_{10} n$ |107309 |181445 |1754047 |12038693 |51523268 |138503720 |428796745 |
| $\sqrt{n}$ | |2440 |60000 |1000000 |16000000 |55000000 |179000000 |
| $n$ |1000000 |60000000 |3600000000 |86400000000 |2592000000000 |31536000000000 |3156000000000000 |
| $n \log_{10} n$ |16404 |162888 |1572184 |10000000 |21858983 |61963474 |185000000 |
| $n^2$ |1000 |2449 |60000|9289 |161602 |179993 |1771012 |
| $n^3$ |100 |394 |1442 |1987 |2954 |3110 |14621 |
| $2^n$ |19 |25 |31 |36 |41 |43 |46 |
| $n!$ |12 |14 |15 |16 |20 |22 |23 |



## Task 1.5

**[2 Point(s)]**

<!-- grade_id: algorithms_introduction_E, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
import numpy as np
time_in_microseconds = {
    "1 second": 1_000_000,
    "1 minute": 60_000_000,
    "1 hour": 3_600_000_000,
    "1 day": 86_400_000_000,
    "1 month": 2_592_000_000_000,
    "1 year": 31_536_000_000_000,
    "1 century": 3_156_000_000_000_000,
}

def log_n():
    results = {}
    for item in time_in_microseconds:
        t = time_in_microseconds[item]
        n = 1  # Start from 1 to avoid log(0)
        while n * np.log(n) <= t:  # Check for log(n) <= t
            n += 1
        results[item] = n - 1  # Store the largest n that satisfies log(n) <= t
    return results

# Print results
print("The results for log(n): {}".format(log_n()))

def sqrt_n():
    results = {}
    for item in time_in_microseconds:
        t = time_in_microseconds[item]
        n = 1  # Start from 1 to avoid sqrt(0)
        while np.sqrt(n) <= t:  # Check for sqrt(n) <= t
            n += 1
        results[item] = n - 1  # Store the largest n that satisfies sqrt(n) <= t
    return results

# Print formatted results
print("The results for sqrt(n): {}".format(sqrt_n()))

def calculate_n():
    results = {}
    for item in time_in_microseconds:
        t = time_in_microseconds[item]
        n = t  # The largest n that can be handled in time t
        results[item] = n  # Store n directly, since n can equal t
    return results

# Print formatted results
print("The results for n: {}".format(calculate_n()))

def n_log_n():
    results = {}
    for item in time_in_microseconds:
        t = time_in_microseconds[item]
        n = 1  # Start from 1 to avoid log(0)
        while n * np.log(n) <= t:  # Check for n * log(n) <= t
            n += 1
        results[item] = n - 1  # Store the largest n that satisfies n * log(n) <= t
    return results

# Print formatted results
print("The results for n * log(n): {}".format(n_log_n()))

def n_squared():
    results = {}
    for item in time_in_microseconds:
        t = time_in_microseconds[item]
        n = 1  # Start from 1
        while n**2 <= t:  # Check for n^2 <= t
            n += 1
        results[item] = n - 1  # Store the largest n that satisfies n^2 <= t
    return results

# Print formatted results
print("The results for n^2: {}".format(n_squared()))

def n_cubed():
    results = {}
    for item in time_in_microseconds:
        t = time_in_microseconds[item]
        n = 1  # Start from 1
        while n**3 <= t:  # Check for n^3 <= t
            n += 1
        results[item] = n - 1  # Store the largest n that satisfies n^3 <= t
    return results

# Print formatted results
print("The results for n^3: {}".format(n_cubed()))

def power_of_two():
    results = {}
    for item in time_in_microseconds:
        t = time_in_microseconds[item]
        n = 0  # Start from 0, since 2^0 = 1
        while 2**n <= t:  # Check for 2^n <= t
            n += 1
        results[item] = n - 1  # Store the largest n that satisfies 2^n <= t
    return results

# Print formatted results
print("The results for 2^n: {}".format(power_of_two()))

def factorial_limit():
    results = {}
    for item in time_in_microseconds:
        t = time_in_microseconds[item]
        n = 0  # Start from 0 since 0! = 1
        factorial = 1  # Initialize factorial(0) = 1
        while factorial <= t:  # Check for n! <= t
            n += 1
            factorial *= n  # Calculate n! incrementally
        results[item] = n - 1  # Store the largest n that satisfies n! <= t
    return results

# Print formatted results
print("The results for n!: {}".format(factorial_limit()))

#raise NotImplementedError()
```

# Task 2

**[80 Point(s)]**


# Uninformed Search

In this Assignment you will implement the unfinformed search strategies *Breadth First Search*, *Depth First Search* and *Iterative Deepening Depth First Search*.

## Task 2.1

**[10 Point(s)]**

## Grid World [10 Points]

You will find three text files, each representing a map for the programming assignment in the maps folder. A search agent is supposed to explore each of these three maps using *breadth-first search*, *depth-first search* and *iterative-deepening depth-first search*.

### Grid World
You will find three text files in `./Complexity_and_Uninformed_Search_files/data`, each representing a map. They are to be interpreted as follows:
- The text file is considered a grid.
- Each character in the text file represents a `cell` in this grid.
  - `*`: goal
  - `space`: free
  - `S`: initial position of the agent
  - Any other character represents an obstacle

**Implement a parsing mechanism to load the `.txt` files into a meaningful data structure into memory. Please also fill the docstring of the function and add typehints such that I can understand what arguments are expected and what it returns.**

<!-- grade_id: uninformed_search_A, points: 0, grade: false, solution: true -->
```python
from typing import List, Tuple

def parse_grid(file_path: str) -> Tuple[List[List[str]], Tuple[int, int], Tuple[int, int]]:
    """
    Parses a grid map from a text file into a 2D list and identifies the start ('S') and goal ('*') positions.

    Parameters
    ----------
    file_path : str
        Path to the text file containing the grid map.

    Returns
    -------
    Tuple[List[List[str]], Tuple[int, int], Tuple[int, int]]
        - A 2D list (grid) representing the map.
        - A tuple (row, col) representing the start position ('S').
        - A tuple (row, col) representing the goal position ('*').

    Raises
    ------
    ValueError
        If the start ('S') or goal ('*') position is not found in the grid.
    """
    grid = []
    start = None
    goal = None

    # Step 1: Read the file without stripping spaces
    try:
        with open(file_path, 'r') as f:
            for line in f:
                grid.append(list(line.rstrip('\n')))  # Only remove newlines, not spaces
    except FileNotFoundError:
        raise ValueError(f"File not found: {file_path}")

    # Step 2: Find the start ('S') and goal ('*') positions
    for i in range(len(grid)):          # Loop through rows
        for j in range(len(grid[i])):  # Loop through columns
            if grid[i][j] == 'S':      # Check for start position
                start = (i, j)
            elif grid[i][j] == '*':    # Check for goal position
                goal = (i, j)

    # Ensure both start and goal positions are found
    if start is None or goal is None:
        raise ValueError("The grid must contain both a start ('S') and a goal ('*').")

    return grid, start, goal


    # YOUR CODE HERE
    #raise NotImplementedError()

# YOUR CODE HERE
#raise NotImplementedError()

```

<!-- grade_id: test_uninformed_search_A0, points: 10, grade: true, solution: false -->
```python
# Test

# Please write your code test here!

```

## Task 2.2

**[10 Point(s)]**

## Agent
The agent must follow the following rules:
- The agent can move from one call to another at each step.
- The agent can **only move to the left, right, up or down** from the current position. The surrounding
cells should be considered as the children of the current cell.
- The agent does **not have any previous knowledge** about the environment (such as goal positions or obstacles) so it has to **explore**.
- The agent **cannot move through obstacles** and the **map is closed**.

## Path Visualization [10 Points]

The function `assign_character` is used to indicate the direction of traversal of the search agent through the nodes of the search tree. Please use the following symbols for visualizing the respective movement of the agent:
 - `U+2500`: ─
 - `U+2502`: │
 - `U+250C`: ┌
 - `U+2510`: ┐
 - `U+2514`: └
 - `U+2518`: ┘
 - `U+251C`: ├
 - `U+2524`: ┤
 - `U+252C`: ┬
 - `U+2534`: ┴
 - `U+253C`: ┼
 - `U+2574`: ╴
 - `U+2575`: ╵
 - `U+2576`: ╶
 - `U+2577`: ╷

With the `write_to_file` function you should be able to save the individual paths to each goal, as well as the final fully explored map with the search path to a `.txt` file.

1. **Implement the `assign_character` function and complete the docstring, such that one can tell how the function is called, what it returns and how it works.**
1. **Implement the `write_to_file` function and comeplete the docstring.**

<!-- grade_id: uninformed_search_B, points: 0, grade: false, solution: true -->
```python
from typing import List, Tuple


class Node:
    """Represents a node in the search tree."""
    def __init__(self, x: int, y: int, parent=None):
        self.x = x
        self.y = y
        self.parent = parent  # Parent node for backtracking


def assign_character(maze_map: List[List[str]], current_node: Node, prev_node: Node):
    """
    Assigns characters to indicate the direction of traversal of the agent.
    
    Parameters
    ----------
    maze_map : List[List[str]]
        The grid map as a 2D list of characters.
    current_node : Node
        The current node in the path.
    prev_node : Node
        The previous node in the path.
    """
    if prev_node is None:
        return  # No character assignment for the start position

    move_x = current_node.x - prev_node.x
    move_y = current_node.y - prev_node.y

    # Determine the character based on movement direction
    if move_x == 1 and move_y == 0:     # Downward move
        maze_map[prev_node.x][prev_node.y] = '│'
    elif move_x == -1 and move_y == 0:  # Upward move
        maze_map[prev_node.x][prev_node.y] = '│'
    elif move_x == 0 and move_y == 1:   # Rightward move
        maze_map[prev_node.x][prev_node.y] = '─'
    elif move_x == 0 and move_y == -1:  # Leftward move
        maze_map[prev_node.x][prev_node.y] = '─'
    else:
        # For advanced intersections and corners, add more cases here if needed
        pass


def path_to_goal(node: Node, maze_map: List[List[str]]):
    """
    Backtracks from the goal node to the start node and assigns path characters.
    
    Parameters
    ----------
    node : Node
        The goal node to start backtracking from.
    maze_map : List[List[str]]
        The grid map as a 2D list.
    """
    while node.parent:  # Until we reach the start node
        assign_character(maze_map, node, node.parent)
        node = node.parent  # Move backward along the path
    maze_map[node.x][node.y] = 'S'  # Mark the start position explicitly


def print_path(maze_map: List[List[str]]):
    """
    Prints the current state of the grid map with the traversal path.
    
    Parameters
    ----------
    maze_map : List[List[str]]
        The grid map as a 2D list.
    """
    for row in maze_map:
        print("".join(row))
    print()


def write_to_file(file_name: str, maze_map: List[List[str]], path: str):
    """
    Writes the traversed path or fully explored map to a `.txt` file.
    
    Parameters
    ----------
    file_name : str
        The name of the output text file.
    maze_map : List[List[str]]
        The grid map as a 2D list, including the traversal path.
    path : str
        The directory where the file will be saved.
    """
    with open(path + file_name, "w", encoding="utf-8") as file:  # Specify UTF-8 encoding
        for row in maze_map:
            file.write("".join(row) + "\n")


def main():
    # Example maze map
    maze_map = [
        ['S', ' ', ' ', ' ', '*', ' '],
        [' ', '#', '#', ' ', '#', ' '],
        [' ', ' ', '#', ' ', '#', ' '],
        [' ', ' ', '#', ' ', ' ', ' '],
        [' ', ' ', ' ', ' ', '#', ' '],
    ]

    # Create nodes for testing
    start_node = Node(0, 0)  # Start position (S)
    goal_node = Node(0, 4, parent=start_node)  # Goal position (*)
    intermediate_node = Node(0, 3, parent=start_node)

    # Simulate a path
    goal_node.parent = intermediate_node
    intermediate_node.parent = start_node

    # Backtrack and mark the path
    path_to_goal(goal_node, maze_map)

    # Print the map with the path
    print("Maze with Path:")
    print_path(maze_map)

    # Save the map to a file
    write_to_file("output_maze.txt", maze_map, "./")


if __name__ == "__main__":
    main()

    # YOUR CODE HERE
    #raise NotImplementedError()
```

<!-- grade_id: test_uninformed_search_B0, points: 10, grade: true, solution: false -->
```python
# Test

# Please write your code test here!

```

## Task 2.3

**[20 Point(s)]**

## Breadth First Search [20 Points]

Implement *Breadth-First Search*. Make sure to fill the docstring of the function and add typehints such that one can understand what arguments are expected and what it returns.

Run the search algorithm for each map and save the results into `./data/results/<search_algo_name>_<map_number>.txt`.

<!-- grade_id: uninformed_search_C, points: 0, grade: false, solution: true -->
```python
import logging
from collections import deque
from typing import List, Tuple

# Set up logging
logging.basicConfig(level=logging.INFO, format='%(asctime)s - %(levelname)s - %(message)s')


class Node:
    """Represents a node in the BFS search tree."""
    def __init__(self, x: int, y: int, state: str = None, parent: 'Node' = None):
        self.x = x
        self.y = y
        self.state = state
        self.parent = parent


def parse_grid(file_path: str) -> Tuple[List[List[str]], Tuple[int, int], Tuple[int, int]]:
    """
    Parses a grid map from a text file into a 2D list and identifies the start ('S') and goal ('*') positions.
    """
    logging.info(f"Parsing the grid from file: {file_path}")
    grid = []
    start = None
    goal = None

    try:
        with open(file_path, 'r') as f:
            for line in f:
                grid.append(list(line.rstrip('\n')))  # Only remove newlines, not spaces
    except FileNotFoundError:
        logging.error(f"File not found: {file_path}")
        raise ValueError(f"File not found: {file_path}")

    for i in range(len(grid)):          # Loop through rows
        for j in range(len(grid[i])):  # Loop through columns
            if grid[i][j].lower() == 's':  # Start position (case-insensitive)
                start = (i, j)
            elif grid[i][j] == '*':    # Goal position
                goal = (i, j)

    if start is None or goal is None:
        logging.error("The grid must contain both a start ('S') and a goal ('*').")
        raise ValueError("The grid must contain both a start ('S') and a goal ('*').")

    logging.info(f"Start position found at: {start}")
    logging.info(f"Goal position found at: {goal}")
    return grid, start, goal


def assign_character(maze_map: List[List[str]], current_node: Node, prev_node: Node):
    """Assigns characters ('─', '│') to indicate the path direction in the maze."""
    move_x = current_node.x - prev_node.x
    move_y = current_node.y - prev_node.y

    if move_x == 1:  # Moving down
        maze_map[prev_node.x][prev_node.y] = '│'
    elif move_x == -1:  # Moving up
        maze_map[prev_node.x][prev_node.y] = '│'
    elif move_y == 1:  # Moving right
        maze_map[prev_node.x][prev_node.y] = '─'
    elif move_y == -1:  # Moving left
        maze_map[prev_node.x][prev_node.y] = '─'


def path_to_goal(node: Node, maze_map: List[List[str]]):
    """Backtracks from the goal node to the start node and assigns path characters."""
    logging.info("Backtracking to mark the path from goal to start.")
    while node.parent:  # Until we reach the start node
        assign_character(maze_map, node, node.parent)
        logging.info(f"Marking path at: ({node.x}, {node.y})")
        node = node.parent  # Move backward along the path
    maze_map[node.x][node.y] = 'S'  # Mark the start position explicitly


def breadth_first_search(grid: List[List[str]], start: Tuple[int, int], goal: Tuple[int, int]) -> List[List[str]]:
    """
    Implements the Breadth-First Search (BFS) algorithm to find the shortest path
    from the start ('S') to the goal ('*') in a grid.
    """
    logging.info("Starting Breadth-First Search (BFS).")
    rows, cols = len(grid), len(grid[0])
    movements = [(-1, 0), (1, 0), (0, -1), (0, 1)]  # Possible moves: up, down, left, right
    queue = deque([Node(start[0], start[1], state='S', parent=None)])  # Initialize the queue with the start node
    visited = set()  # Set to keep track of visited cells

    while queue:
        node = queue.popleft()  # Take the first node from the queue
        visited.add((node.x, node.y))  # Mark it as visited
        logging.info(f"Visiting node at: ({node.x}, {node.y})")

        # Check if the goal is reached
        if (node.x, node.y) == goal:
            logging.info(f"Goal reached at: ({node.x}, {node.y})")
            path_to_goal(node, grid)  # Backtrack to mark the path on the grid
            return grid  # Return the updated grid

        # Explore neighbors
        for move_x, move_y in movements:
            new_x, new_y = node.x + move_x, node.y + move_y

            # Check if the new position is within bounds, not visited, and passable
            if (
                0 <= new_x < rows and 0 <= new_y < cols and
                (new_x, new_y) not in visited and
                grid[new_x][new_y] in (' ', '*')
            ):
                logging.info(f"Adding neighbor to queue: ({new_x}, {new_y})")
                queue.append(Node(new_x, new_y, state=grid[new_x][new_y], parent=node))

    logging.error("No path found from start to goal.")
    raise ValueError("No path found from start to goal.")


def save_grid_to_file(file_path: str, grid: List[List[str]]):
    """Saves the grid to a text file."""
    logging.info(f"Saving the grid to file: {file_path}")
    with open(file_path, 'w') as f:
        for row in grid:
            f.write(''.join(row) + '\n')


def main():
    input_file = "map1.txt"  # Path to your maze file
    output_file = "output_maze.txt"  # Path to save the output

    try:
        grid, start, goal = parse_grid(input_file)
        marked_grid = breadth_first_search(grid, start, goal)

        logging.info("Maze with path:")
        for row in marked_grid:
            logging.info("".join(row))

        save_grid_to_file(output_file, marked_grid)
        logging.info(f"Result saved to {output_file}")

    except ValueError as e:
        logging.error(e)


if __name__ == "__main__":
    main()

```

```python
import logging
from typing import List, Tuple
from collections import deque

# Set up logging
logging.basicConfig(level=logging.INFO, format='%(asctime)s - %(levelname)s - %(message)s')


class Node:
    """Represents a node in the BFS search tree."""
    def __init__(self, x: int, y: int, parent: 'Node' = None):
        self.x = x
        self.y = y
        self.parent = parent


def parse_grid(file_path: str) -> Tuple[List[List[str]], Tuple[int, int], Tuple[int, int]]:
    """
    Parses a grid map from a text file into a 2D list and identifies the start ('S') and goal ('*') positions.
    """
    logging.info(f"Parsing the grid from file: {file_path}")
    grid = []
    start = None
    goal = None

    try:
        with open(file_path, 'r') as f:
            for line in f:
                grid.append(list(line.rstrip('\n')))  # Only remove newlines, not spaces
    except FileNotFoundError:
        logging.error(f"File not found: {file_path}")
        raise ValueError(f"File not found: {file_path}")

    for i in range(len(grid)):          # Loop through rows
        for j in range(len(grid[i])):  # Loop through columns
            if grid[i][j].lower() == 's':  # Start position (case-insensitive)
                start = (i, j)
            elif grid[i][j] == '*':    # Goal position
                goal = (i, j)

    if start is None or goal is None:
        logging.error("The grid must contain both a start ('S') and a goal ('*').")
        raise ValueError("The grid must contain both a start ('S') and a goal ('*').")

    logging.info(f"Start position found at: {start}")
    logging.info(f"Goal position found at: {goal}")
    return grid, start, goal


def path_to_goal(node: Node, maze_map: List[List[str]]):
    """Backtracks from the goal node to the start node and assigns path characters."""
    logging.info("Backtracking to mark the path from goal to start.")
    while node.parent:  # Until we reach the start node
        prev_node = node.parent
        move_x = node.x - prev_node.x
        move_y = node.y - prev_node.y

        # Assign path characters based on the direction
        if move_x == 1 or move_x == -1:  # Vertical movement
            maze_map[prev_node.x][prev_node.y] = '│'
        elif move_y == 1 or move_y == -1:  # Horizontal movement
            maze_map[prev_node.x][prev_node.y] = '─'

        logging.info(f"Marking path at: ({node.x}, {node.y})")
        node = prev_node  # Move backward along the path

    maze_map[node.x][node.y] = 'S'  # Mark the start position explicitly


def breadth_first_search(grid: List[List[str]], start: Tuple[int, int], goal: Tuple[int, int]) -> List[List[str]]:
    """
    Implements the Breadth-First Search (BFS) algorithm to find the shortest path
    from the start ('S') to the goal ('*') in a grid.

    Parameters
    ----------
    grid : List[List[str]]
        A 2D list representing the grid.
    start : Tuple[int, int]
        The (row, col) coordinates of the start position ('S').
    goal : Tuple[int, int]
        The (row, col) coordinates of the goal position ('*').

    Returns
    -------
    List[List[str]]
        The grid with the shortest path marked from the start ('S') to the goal ('*').
    """
    logging.info("Starting Breadth-First Search (BFS).")
    rows, cols = len(grid), len(grid[0])
    movements = [(-1, 0), (1, 0), (0, -1), (0, 1)]  # Possible moves: up, down, left, right
    visited = set()  # Set to keep track of visited cells
    queue = deque([Node(start[0], start[1])])  # Initialize BFS queue with the start node

    while queue:
        current_node = queue.popleft()
        x, y = current_node.x, current_node.y

        # Check if we reached the goal
        if (x, y) == goal:
            logging.info(f"Goal reached at: ({x}, {y})")
            path_to_goal(Node(x=x, y=y, parent=current_node.parent), grid)
            return grid

        # Mark the current node as visited
        visited.add((x, y))
        logging.info(f"Visiting node at: ({x}, {y})")

        # Explore neighbors
        for move_x, move_y in movements:
            new_x, new_y = x + move_x, y + move_y
            if (
                0 <= new_x < rows and 0 <= new_y < cols and
                grid[new_x][new_y] in (' ', '*') and
                (new_x, new_y) not in visited
            ):
                queue.append(Node(new_x, new_y, parent=current_node))

    # If we finish the loop and didn't find a path
    logging.error("No path found from start to goal.")
    raise ValueError("No path found from start to goal.")


def save_grid_to_file(file_path: str, grid: List[List[str]]):
    """Saves the grid to a text file."""
    logging.info(f"Saving the grid to file: {file_path}")
    try:
        with open(file_path, 'w', encoding='utf-8') as f:  # Use UTF-8 encoding
            for row in grid:
                f.write(''.join(row) + '\n')
        logging.info(f"Grid successfully saved to {file_path}")
    except Exception as e:
        logging.error(f"Error saving grid to file: {e}")


def main():
    input_file = "map1.txt"  # Path to your maze file
    output_file = "output_shortest_path.txt"  # Path to save the shortest path output

    try:
        grid, start, goal = parse_grid(input_file)
        marked_grid = breadth_first_search(grid, start, goal)

        logging.info("Maze with shortest path:")
        for row in marked_grid:
            logging.info("".join(row))

        save_grid_to_file(output_file, marked_grid)
        logging.info(f"Result saved to {output_file}")

    except ValueError as e:
        logging.error(e)


if __name__ == "__main__":
    main()

```

<!-- grade_id: test_uninformed_search_C0, points: 20, grade: true, solution: false -->
```python
# Test

# Please write your code test here!

```

## Task 2.4

**[20 Point(s)]**

## Depth First Search [20 Points]

Implement *Depth-First Search*. Make sure to fill the docstring of the function and add typehints such that one can understand what arguments are expected and what it returns.

Run the search algorithm for each map and save the results into `./data/results/<search_algo_name>_<map_number>.txt`.

<!-- grade_id: uninformed_search_D, points: 0, grade: false, solution: true -->
```python
import logging
from typing import List, Tuple

# Set up logging
logging.basicConfig(level=logging.INFO, format='%(asctime)s - %(levelname)s - %(message)s')


class Node:
    """Represents a node in the DFS search tree."""
    def __init__(self, x: int, y: int, state: str = None, parent: 'Node' = None):
        self.x = x
        self.y = y
        self.state = state
        self.parent = parent


def parse_grid(file_path: str) -> Tuple[List[List[str]], Tuple[int, int], Tuple[int, int]]:
    """
    Parses a grid map from a text file into a 2D list and identifies the start ('S') and goal ('*') positions.
    """
    logging.info(f"Parsing the grid from file: {file_path}")
    grid = []
    start = None
    goal = None

    try:
        with open(file_path, 'r') as f:
            for line in f:
                grid.append(list(line.rstrip('\n')))  # Only remove newlines, not spaces
    except FileNotFoundError:
        logging.error(f"File not found: {file_path}")
        raise ValueError(f"File not found: {file_path}")

    for i in range(len(grid)):          # Loop through rows
        for j in range(len(grid[i])):  # Loop through columns
            if grid[i][j].lower() == 's':  # Start position (case-insensitive)
                start = (i, j)
            elif grid[i][j] == '*':    # Goal position
                goal = (i, j)

    if start is None or goal is None:
        logging.error("The grid must contain both a start ('S') and a goal ('*').")
        raise ValueError("The grid must contain both a start ('S') and a goal ('*').")

    logging.info(f"Start position found at: {start}")
    logging.info(f"Goal position found at: {goal}")
    return grid, start, goal


def assign_character(maze_map: List[List[str]], current_node: Node, prev_node: Node):
    """Assigns characters ('─', '│') to indicate the path direction in the maze."""
    move_x = current_node.x - prev_node.x
    move_y = current_node.y - prev_node.y

    if move_x == 1:  # Moving down
        maze_map[prev_node.x][prev_node.y] = '│'
    elif move_x == -1:  # Moving up
        maze_map[prev_node.x][prev_node.y] = '│'
    elif move_y == 1:  # Moving right
        maze_map[prev_node.x][prev_node.y] = '─'
    elif move_y == -1:  # Moving left
        maze_map[prev_node.x][prev_node.y] = '─'


def path_to_goal(node: Node, maze_map: List[List[str]]):
    """Backtracks from the goal node to the start node and assigns path characters."""
    logging.info("Backtracking to mark the path from goal to start.")
    while node.parent:  # Until we reach the start node
        assign_character(maze_map, node, node.parent)
        logging.info(f"Marking path at: ({node.x}, {node.y})")
        node = node.parent  # Move backward along the path
    maze_map[node.x][node.y] = 'S'  # Mark the start position explicitly


def depth_first_search(grid: List[List[str]], start: Tuple[int, int], goal: Tuple[int, int]) -> List[List[str]]:
    """
    Implements the Depth-First Search (DFS) algorithm to find a path
    from the start ('S') to the goal ('*') in a grid.

    Parameters
    ----------
    grid : List[List[str]]
        A 2D list representing the grid.
    start : Tuple[int, int]
        The (row, col) coordinates of the start position ('S').
    goal : Tuple[int, int]
        The (row, col) coordinates of the goal position ('*').

    Returns
    -------
    List[List[str]]
        The grid with the path marked from the start ('S') to the goal ('*').
    """
    logging.info("Starting Depth-First Search (DFS).")
    rows, cols = len(grid), len(grid[0])
    movements = [(-1, 0), (1, 0), (0, -1), (0, 1)]  # Possible moves: up, down, left, right
    visited = set()  # Set to keep track of visited cells

    def dfs_recursive(x: int, y: int, parent: Node) -> bool:
        """Recursive helper function for DFS."""
        if (x, y) in visited:
            return False  # Already visited

        visited.add((x, y))
        logging.info(f"Visiting node at: ({x}, {y})")

        # Check if the goal is reached
        if (x, y) == goal:
            logging.info(f"Goal reached at: ({x}, {y})")
            path_to_goal(Node(x=x, y=y, state='*', parent=parent), grid)
            return True

        # Explore neighbors
        for move_x, move_y in movements:
            new_x, new_y = x + move_x, y + move_y
            if (
                0 <= new_x < rows and 0 <= new_y < cols and
                grid[new_x][new_y] in (' ', '*')
            ):
                if dfs_recursive(new_x, new_y, Node(x=x, y=y, state=grid[x][y], parent=parent)):
                    return True  # Path found

        return False  # No path found from this node

    # Start DFS from the starting position
    if not dfs_recursive(start[0], start[1], None):
        logging.error("No path found from start to goal.")
        raise ValueError("No path found from start to goal.")

    return grid


def save_grid_to_file(file_path: str, grid: List[List[str]]):
    """Saves the grid to a text file."""
    logging.info(f"Saving the grid to file: {file_path}")
    try:
        with open(file_path, 'w', encoding='utf-8') as f:  # Use UTF-8 encoding
            for row in grid:
                f.write(''.join(row) + '\n')
        logging.info(f"Grid successfully saved to {file_path}")
    except Exception as e:
        logging.error(f"Error saving grid to file: {e}")


def main():
    input_file = "map1.txt"  # Path to your maze file
    output_file = "output_maze_dfs.txt"  # Path to save the output

    try:
        grid, start, goal = parse_grid(input_file)
        marked_grid = depth_first_search(grid, start, goal)

        logging.info("Maze with path:")
        for row in marked_grid:
            logging.info("".join(row))

        save_grid_to_file(output_file, marked_grid)
        logging.info(f"Result saved to {output_file}")

    except ValueError as e:
        logging.error(e)


if __name__ == "__main__":
    main()

```

Output:
```
2024-11-24 00:22:00,101 - INFO - Parsing the grid from file: map1.txt
2024-11-24 00:22:00,103 - INFO - Start position found at: (1, 3)
2024-11-24 00:22:00,104 - INFO - Goal position found at: (19, 72)
2024-11-24 00:22:00,104 - INFO - Starting Depth-First Search (DFS).
2024-11-24 00:22:00,105 - INFO - Visiting node at: (1, 3)
2024-11-24 00:22:00,106 - INFO - Visiting node at: (2, 3)
2024-11-24 00:22:00,107 - INFO - Visiting node at: (2, 2)
2024-11-24 00:22:00,108 - INFO - Visiting node at: (1, 2)
2024-11-24 00:22:00,108 - INFO - Visiting node at: (1, 1)
2024-11-24 00:22:00,108 - INFO - Visiting node at: (2, 1)
2024-11-24 00:22:00,108 - INFO - Visiting node at: (2, 4)
2024-11-24 00:22:00,110 - INFO - Visiting node at: (1, 4)
2024-11-24 00:22:00,110 - INFO - Visiting node at: (1, 5)
2024-11-24 00:22:00,111 - INFO - Visiting node at: (2, 5)
2024-11-24 00:22:00,111 - INFO - Visiting node at: (2, 6)
2024-11-24 00:22:00,111 - INFO - Visiting node at: (1, 6)
2024-11-24 00:22:00,112 - INFO - Vi
```

```python

```

```python
import logging
from typing import List, Tuple

# Set up logging
logging.basicConfig(level=logging.INFO, format='%(asctime)s - %(levelname)s - %(message)s')


class Node:
    """Represents a node in the DFS search tree."""
    def __init__(self, x: int, y: int, parent: 'Node' = None):
        self.x = x
        self.y = y
        self.parent = parent


def parse_grid(file_path: str) -> Tuple[List[List[str]], Tuple[int, int], Tuple[int, int]]:
    """
    Parses a grid map from a text file into a 2D list and identifies the start ('S') and goal ('*') positions.
    """
    logging.info(f"Parsing the grid from file: {file_path}")
    grid = []
    start = None
    goal = None

    try:
        with open(file_path, 'r') as f:
            for line in f:
                grid.append(list(line.rstrip('\n')))  # Only remove newlines, not spaces
    except FileNotFoundError:
        logging.error(f"File not found: {file_path}")
        raise ValueError(f"File not found: {file_path}")

    for i in range(len(grid)):          # Loop through rows
        for j in range(len(grid[i])):  # Loop through columns
            if grid[i][j].lower() == 's':  # Start position (case-insensitive)
                start = (i, j)
            elif grid[i][j] == '*':    # Goal position
                goal = (i, j)

    if start is None or goal is None:
        logging.error("The grid must contain both a start ('S') and a goal ('*').")
        raise ValueError("The grid must contain both a start ('S') and a goal ('*').")

    logging.info(f"Start position found at: {start}")
    logging.info(f"Goal position found at: {goal}")
    return grid, start, goal


def path_to_goal(node: Node, maze_map: List[List[str]]):
    """Backtracks from the goal node to the start node and assigns path characters."""
    logging.info("Backtracking to mark the path from goal to start.")
    while node.parent:  # Until we reach the start node
        prev_node = node.parent
        move_x = node.x - prev_node.x
        move_y = node.y - prev_node.y

        # Assign path characters based on the direction
        if move_x == 1 or move_x == -1:  # Vertical movement
            maze_map[prev_node.x][prev_node.y] = '│'
        elif move_y == 1 or move_y == -1:  # Horizontal movement
            maze_map[prev_node.x][prev_node.y] = '─'

        logging.info(f"Marking path at: ({node.x}, {node.y})")
        node = prev_node  # Move backward along the path

    maze_map[node.x][node.y] = 'S'  # Mark the start position explicitly


def depth_first_search(grid: List[List[str]], start: Tuple[int, int], goal: Tuple[int, int]) -> List[List[str]]:
    """
    Implements the Depth-First Search (DFS) algorithm to find a path
    from the start ('S') to the goal ('*') in a grid.

    Parameters
    ----------
    grid : List[List[str]]
        A 2D list representing the grid.
    start : Tuple[int, int]
        The (row, col) coordinates of the start position ('S').
    goal : Tuple[int, int]
        The (row, col) coordinates of the goal position ('*').

    Returns
    -------
    List[List[str]]
        The grid with the path marked from the start ('S') to the goal ('*').
    """
    logging.info("Starting Depth-First Search (DFS).")
    rows, cols = len(grid), len(grid[0])
    movements = [(-1, 0), (1, 0), (0, -1), (0, 1)]  # Possible moves: up, down, left, right
    visited = set()  # Set to keep track of visited cells

    def dfs_recursive(x: int, y: int, parent: Node) -> bool:
        """Recursive helper function for DFS."""
        if (x, y) in visited:
            return False  # Already visited

        visited.add((x, y))
        logging.info(f"Visiting node at: ({x}, {y})")

        # Check if the goal is reached
        if (x, y) == goal:
            logging.info(f"Goal reached at: ({x}, {y})")
            path_to_goal(Node(x=x, y=y, parent=parent), grid)
            return True

        # Explore neighbors
        for move_x, move_y in movements:
            new_x, new_y = x + move_x, y + move_y
            if (
                0 <= new_x < rows and 0 <= new_y < cols and
                grid[new_x][new_y] in (' ', '*') and
                (new_x, new_y) not in visited
            ):
                if dfs_recursive(new_x, new_y, Node(x=x, y=y, parent=parent)):
                    return True  # Path found

        return False  # No path found from this node

    # Start DFS from the starting position
    if not dfs_recursive(start[0], start[1], None):
        logging.error("No path found from start to goal.")
        raise ValueError("No path found from start to goal.")

    return grid


def save_grid_to_file(file_path: str, grid: List[List[str]]):
    """Saves the grid to a text file."""
    logging.info(f"Saving the grid to file: {file_path}")
    try:
        with open(file_path, 'w', encoding='utf-8') as f:  # Use UTF-8 encoding
            for row in grid:
                f.write(''.join(row) + '\n')
        logging.info(f"Grid successfully saved to {file_path}")
    except Exception as e:
        logging.error(f"Error saving grid to file: {e}")


def main():
    input_file = "map1.txt"  # Path to your maze file
    output_file = "output_dfs_shortest_path.txt"  # Path to save the output

    try:
        grid, start, goal = parse_grid(input_file)
        marked_grid = depth_first_search(grid, start, goal)

        logging.info("Maze with path:")
        for row in marked_grid:
            logging.info("".join(row))

        save_grid_to_file(output_file, marked_grid)
        logging.info(f"Result saved to {output_file}")

    except ValueError as e:
        logging.error(e)


if __name__ == "__main__":
    main()

```

Output:
```
2024-11-24 00:33:47,193 - INFO - Parsing the grid from file: map1.txt
2024-11-24 00:33:47,207 - INFO - Start position found at: (1, 3)
2024-11-24 00:33:47,208 - INFO - Goal position found at: (19, 72)
2024-11-24 00:33:47,208 - INFO - Starting Depth-First Search (DFS).
2024-11-24 00:33:47,209 - INFO - Visiting node at: (1, 3)
2024-11-24 00:33:47,210 - INFO - Visiting node at: (2, 3)
2024-11-24 00:33:47,211 - INFO - Visiting node at: (2, 2)
2024-11-24 00:33:47,211 - INFO - Visiting node at: (1, 2)
2024-11-24 00:33:47,211 - INFO - Visiting node at: (1, 1)
2024-11-24 00:33:47,212 - INFO - Visiting node at: (2, 1)
2024-11-24 00:33:47,213 - INFO - Visiting node at: (2, 4)
2024-11-24 00:33:47,213 - INFO - Visiting node at: (1, 4)
2024-11-24 00:33:47,213 - INFO - Visiting node at: (1, 5)
2024-11-24 00:33:47,214 - INFO - Visiting node at: (2, 5)
2024-11-24 00:33:47,215 - INFO - Visiting node at: (2, 6)
2024-11-24 00:33:47,215 - INFO - Visiting node at: (1, 6)
2024-11-24 00:33:47,216 - INFO - Vi
```

<!-- grade_id: test_uninformed_search_D0, points: 20, grade: true, solution: false -->
```python
# Test

# Please write your code test here!

```

## Task 2.5

**[20 Point(s)]**

## Iterative Deepening Search [20 Points]

Implement *Iterative Deepening Search*. Make sure to fill the docstring of the function and add typehints such that one can understand what arguments are expected and what it returns.

Run the search algorithm for each map and save the results into `./data/results/<search_algo_name>_<map_number>.txt`.

<!-- grade_id: uninformed_search_E, points: 0, grade: false, solution: true -->
```python
def iterative_deepening_depth_first_search(maze_map, monitor=False):
    """ This function implements the iterative deepening depth-first search (IDDFS) algorithm.
    THe key point about IDDFS is that it uses depth-limited depth-first search with an incrementing depth
    limit and thereby explores the node basically breadth first, but at the same time makes use of the space
    saving procedure of DFS, which keeps only b number of nodes in the memory, with b being the branching
    factor of the current node.
    
    Parameters
    ----------
    maze_map : [type]
        [description]
    monitor : bool
        Flag for enabling or disabling command line output for the search run. Defaults to False: No
        monitoring.
    
    Returns
    -------
    [type]
        [description]
    """

    
    

    # YOUR CODE HERE
    #raise NotImplementedError()
```

<!-- grade_id: test_uninformed_search_E0, points: 20, grade: true, solution: false -->
```python
# Test

# Please write your code test here!

```